<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Часть-1---Параметрические-гипотезы" data-toc-modified-id="Часть-1---Параметрические-гипотезы-1">Часть 1 - Параметрические гипотезы</a></span><ul class="toc-item"><li><span><a href="#Мудрость" data-toc-modified-id="Мудрость-1.1">Мудрость</a></span></li><li><span><a href="#Тайлер" data-toc-modified-id="Тайлер-1.2">Тайлер</a></span></li><li><span><a href="#Рейтинги-президентов" data-toc-modified-id="Рейтинги-президентов-1.3">Рейтинги президентов</a></span></li><li><span><a href="#Сон" data-toc-modified-id="Сон-1.4">Сон</a></span></li><li><span><a href="#Say-my-name" data-toc-modified-id="Say-my-name-1.5">Say my name</a></span></li><li><span><a href="#Матан" data-toc-modified-id="Матан-1.6">Матан</a></span></li><li><span><a href="#Бутерброд" data-toc-modified-id="Бутерброд-1.7">Бутерброд</a></span></li></ul></li><li><span><a href="#Часть-2---Непараметрические-гипотезы" data-toc-modified-id="Часть-2---Непараметрические-гипотезы-2">Часть 2 - Непараметрические гипотезы</a></span><ul class="toc-item"><li><span><a href="#Однородная-игра" data-toc-modified-id="Однородная-игра-2.1">Однородная игра</a></span></li><li><span><a href="#Мыши-и-зеркало" data-toc-modified-id="Мыши-и-зеркало-2.2">Мыши и зеркало</a></span></li><li><span><a href="#Энергетик" data-toc-modified-id="Энергетик-2.3">Энергетик</a></span></li><li><span><a href="#Бомбардировки" data-toc-modified-id="Бомбардировки-2.4">Бомбардировки</a></span></li></ul></li><li><span><a href="#Часть-3---Самостоятельный-анализ" data-toc-modified-id="Часть-3---Самостоятельный-анализ-3">Часть 3 - Самостоятельный анализ</a></span><ul class="toc-item"><li><span><a href="#Анорексия" data-toc-modified-id="Анорексия-3.1">Анорексия</a></span></li><li><span><a href="#Метилфенидат" data-toc-modified-id="Метилфенидат-3.2">Метилфенидат</a></span></li><li><span><a href="#Сиэтл" data-toc-modified-id="Сиэтл-3.3">Сиэтл</a></span></li><li><span><a href="#Расход-топлива" data-toc-modified-id="Расход-топлива-3.4">Расход топлива</a></span><ul class="toc-item"><li><ul class="toc-item"><li><span><a href="#vs,-mpg" data-toc-modified-id="vs,-mpg-3.4.0.1">vs, mpg</a></span></li><li><span><a href="#am,-mpg" data-toc-modified-id="am,-mpg-3.4.0.2">am, mpg</a></span></li><li><span><a href="#vs,-am" data-toc-modified-id="vs,-am-3.4.0.3">vs, am</a></span></li></ul></li></ul></li></ul></li></ul></div>

В этом домашнем задании предстоит проверить много гипотез.

Каждая гипотеза дает от 5 до 10 баллов и все они поделены на 3 части. Ориентир по заданиям ниже.


Все датасеты лежат [тут](https://drive.google.com/open?id=1-psTZYeGvnr7UKt77X3D271N5BzNRxzZ). Можно скачать эту папку, положить в соседнюю с рабочими ноутбуками и указать соответствующий путь `'../data/something.csv'`

В классрум нужно сдать заполненный ipynb ноутбук с кодом и комментариями. Просьба назвать файлы в формате: `Иванова.ipynb`. Архивировать файл и присылать датасеты НЕ нужно.


**Часть 1 - Параметрические гипотезы**

| Задание              | Баллы |
|----------------------|-------|
| Мудрость             |     5 |
| Тайлер               |     5 |
| Рейтинги президентов |     5 |
| Сон                  |     5 |
| Say my name          |     5 |
| Матан                |     5 |
| Бутерброд            |     5 |
| Всего                |    35 |


**Часть 2 - Непараметрические гипотезы**

| Задание            | Датасет         | Баллы |
|--------------------|-----------------|-------|
| Однородная игра    | killer_time     |     5 |
| Мыши и зеркало     | mouses_mirror   |     5 |
| Энергетик          | energy drink    |     5 |
| Бомбардировки      | london          |     10 |
| | Всего                |    25 |

**Часть 3 - Самостоятельный анализ**

| Задание            | Датасет         | Баллы |
|--------------------|-----------------|-------|
| Анорексия          | anorexia        |     10 |
| Метилфенидат       | methylphenidate |     5 |
| Сиэтл              | seattle         |     5 |
| Расход топлива - 1 | mtcars: vs mpg  |     5 |
| Расход топлива - 2 | mtcars: am mpg  |     5 |
| Расход топлива - 3 | mtcars: am vs   |     10 |
|                    | Всего           |    40 |


# Часть 1 - Параметрические гипотезы


```python
import numpy as np
import pandas as pd

from scipy import stats as sts

import seaborn as sns
import matplotlib.pyplot as plt

plt.style.use('ggplot')
```

## Мудрость

> Помните, что для человека звук его имени является самым сладким и самым важным звуком человеческой речи. (Дейл Карнеги)

Доверяй, но проверяй. Решим задачку по психологии! Все цитаты в пабликах вконтакте делятся на два вида: мудрые и недостаточно мудрые. Если цитата мудрая, то её можно лайкнуть. Поговаривают, что если в цитате есть личные местоимения, то пользователь автоматически сопоставляет себя автору этой цитаты и вероятность того, что она будет лайкнута, повышается. Например, цитату 

> Можно простить человеку все, кроме отсутствия. (Экзюпери) 

мне не очень хотелось бы лайкать. А вот цитата 

> Мне наплевать, что вы обо мне думаете. Я о вас вообще не думаю. (Булгаков) 

совершенно другое дело. 

В табличке `df_vk_wise` лежат данные о мудрых цитатах из [мудрого паблика.](https://vk.com/ooioioioioioo) В колонке `I` находится переменная, которая принимает значение `True`, когда в тексте цитаты было личное местоимение "Я".

Срежем первые $50$ постов, так как под ними количество репостов и лайков ещё окончательно не выставлено, и будем работать с остатками. 

```python
df = pd.read_csv('./data/df_vk_wise.csv')[50:]
df.head()
```

Постройте гистограмму для числа лайков постов, содержащих личные местоимения и не содержащих. Как думаете, есть ли с распределением лайков какие-то проблемы?

<!-- grade_id: cell-5b1acf05ccd983b3, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

```

Да, проблемы есть. Это выбросы. Хвосты на гистограммах оказываются очень длинными. Найдите медианы и средние для числа лайков. Сравните их между собой. 

<!-- grade_id: cell-84afb85a40373b11, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

```

Проверьте гипотезу о равенстве среднего числа лайков. Используйте для этого $z$-тест, основанный на ЦПТ. Дисперсию расчитывайте по выборке, используя несмещённую оценку `(ddof=1)`. Запишите наблюдаемое значение статистики в переменную `z_obs`. 

<!-- grade_id: cell-1acfb89f07181a0a, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

alpha = 0.05 
z_obs = ...


```

Найдите критическое значение статистики. Сравните его с наблюдаемым. Рассчитайте p-value для ситуации с двусторонней альтернативой. Сделайте вывод о выдвинутой гипотезе. Критическое значение статистики запишите в переменную `z_crit`, а p-значение в переменную `pval`. 

<!-- grade_id: cell-9ebff976022e649a, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below
z_crit = ...
pval = ...

```

```python
# проверка того, что задание решено верно
assert np.abs(z_obs - 2.6644610922505323) < 1e-5
```

Можно заметить, что у распределения числа репостов длинные хвосты. В данных есть выбросы, которые завышают среднее значение. Если выбросы сильные, то ЦПТ может ломаться. Из-за этого имеет смысл в данном случае проверять гипотезу о равенстве медиан, которые устойчивы к выбросам.

-------------------

## Тайлер

__Люди всё время меня спрашивают, знаю ли я Тайлера Дёрдена?__ Да, знаю. Он делает лучшее мыло в стране. 

__а)__ Тайлер считает партию изделий хорошей, если дисперсия размера не превышает $25$ грамм$^2$. Пока Тайлер летел в самолёте, он измерил $10$ кусков мыла. Выборочная дисперсия оказалась равна $27$ грамм$^2$. Можно ли продавать эту партию или её надо списать в брак? 

При проверке этой гипотезы найдите $p$-значение и запишите его в переменную `pval1`. Наблюдаемое значение статистики запишите в переменную `chi_obs`. Используйте уровень значимости $1\%$.

<!-- grade_id: cell-2e30e820f90148d5, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

chi_obs = ...


```

__б)__ Для того, чтобы уменьшить количество бракованного мыла, Тайлер собирается внедрить новую технологию производства. Тайлер произвёл по старой технологии $100$ кусков мыла. Стандартное отклонение этой партии составило $6$ грамм. Для  новой технологии стандартное отклонение, посчитанное по $100$ кускам мыла, составило $5$ грамм. Правда ли, что новая технология обладает меньшей дисперсией и её нужно внедрить на производстве?

При проверке этой гипотезы найдите $p$-значение и запишите его в переменную `pval2`. Используйте уровень значимости $1\%$.

<!-- grade_id: cell-88a96be4941d7e62, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

pval2 = ...


```

```python
# проверка того, что задание решено верно
assert np.abs(chi_obs - 9.72) < 1e-1
assert np.abs(F_obs - 1.44) < 1e-1

```

## Рейтинги президентов

В таблице `presidents.csv` представлены рейтинги различных президентов США. Социальные опросы проводились раз в квартал. Опрашивалась случайная выборка граждан. У каждого спрашивалось одобряет ли он деятельность президента.

__Люди всё время меня спрашивают, знаю ли я Никсона.__ Да, Ричард Никсон был президентом США с 20 января 1969 года по 9 августа 1974 года. 17 июня 1972 года произошёл Уотергейтский скандал. Давайте проанализируем, что при этом произошло с его рейтингом.

```python
df = pd.read_csv('./data/presidents.csv')
df.head()
```

Чтобы привести таблицу к удобному виду используйте функцию `pd.melt`. Вытащите из таблички рейтинг Никсона с 1 квартала 1969 года по 2 квартал 1972 года (включая его). Запишите эти данные в переменную `x`. Рейтинг с 3 квартала 1972 года по 2 квартал 1974 года поместите в переменную `y`. Пропуски в данных удалите из выборки.

<!-- grade_id: cell-0c35f4bdddbb0029, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

x = ...
y = ... 
 
```

Правда ли, что после скандала средний рейтинг Никсона рухнул? Предполагайте, что рейтинг распределён нормально, а также, что дисперсия рейтинга не изменилась из-за скандала. Наблюдаемое значение статистики запишите в переменную `obs`, в переменную `pval` запишите $p$-значение.

<!-- grade_id: cell-468606ed00795265, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

obs = ...
pval = ...

```

```python
# проверка того, что задание решено верно
assert np.abs(obs - 3.5640741313229793) < 1e-5

```

Снова президентом Никсон станет в $3000$ году. 

![](https://vignette.wikia.nocookie.net/futurama/images/2/2f/Richard_M._Nixon%27s_head.jpg/revision/latest/scale-to-width-down/250?cb=20131123190522&path-prefix=ru)

## Сон

Проводится эксперимент о пользе нового снотворного. Для его проведения используется группа из $10$ студентов. В первый день каждый из них ложится спать. На утро измеряется, сколько он проспал. Из результата вычитают $8$ (нормальное число часов сна). На второй день каждому дают таблетку. Утром снова измеряют сколько человек проспал, и вычитают $8$. Результаты эксперимента записаны в табличку `sleep.csv`.

```python
df = pd.read_csv('./data/sleep.csv')
df.head()
```

Правда ли, что сон студентов до приёма снотворного не отличается от нормы (отклонение от 8 часов равно нулю). Проверьте эту гипотезу на $5\%$ уровне значимости, в переменную `pval1` запишите p-значение. 

<!-- grade_id: cell-c04ca7c5b2d46f59, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

tobs1 = ...
pval1 = ...


```

Правда ли, что благодаря таблеткам студенты стали спать больше (односторонняя альтернатива)? Запишите в переменную `pval2` соотвествующее p-значение. 

<!-- grade_id: cell-4c7c2e8e8b3c5de8, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

tobs2 = ...
pval2 = ...

# your code here

```

<!-- grade_id: cell-2167298693b77a31, points: 1, grade: true, solution: false -->
```python
# проверка того, что задание решено верно
assert np.abs(tobs1 - 1.325710) < 1e-5
assert np.abs(tobs2 - 4.062127) < 1e-5

```

## Say my name

__Люди всё время меня спрашивают, знает ли меня Барсик.__  Многие владельцы кошек уверены, что их питомцы могут понимать достаточно большое количество слов, а также узнавать собственное имя. Однако долгое время никто не пытался доказать это экспериментально. Разобраться с этой проблемой [решили учёные из Токийского университета.](https://nplus1.ru/news/2019/04/04/cat-name?utm_source=obrazovach&utm_medium=social&utm_campaign=yaponskie-uchenye-dokazali--chto-domashnie-k)

Было проведено четыре эксперимента с участием домашних кошек и кошек из кошачьих кафе. Учёные использовали метод привыкания-отвыкания (habituation-dishabituation), который часто используется в экспериментах с участием маленьких детей. 

В соответствии с гипотезой, которая лежит в основе этого метода, определить реакцию на какой-либо стимул можно, представив его после схожих с ним стимулов. К примеру, для того, чтобы выяснить, может ли младенец узнать свое имя, ему представляют несколько слов подряд: с каждым новым словом его реакция снижается, но появление имени должно (в случае, если ребенок действительно узнает свое имя) ее вернуть. Этот же метод применили и к кошкам: реакция кошек на слова характеризовалась поворотом головы, движением ушей, вокализацией, движениями хвоста или перемещениями.

В первом эксперименте исследователи проследили за реакцией кошек на собственное имя, произнесенное хозяином, после четырех существительных (16 кошек, живущих дома с двумя другими кошками или меньше). Данные об этом эксперименте находятся в табличке `cats_names.csv`. 

```python
df = pd.read_csv('./data/cats_names.csv', sep='\t')
df.head()
```

Будем проверять гипотезу о том, что кошки понимают своё имя с помощью упрощённой процедуры: 

- Вытащим из таблички все строчки, где экспериментаторы называли $4-$ое существительное. Они помечены в колонке `stimuli` как `Noun 4`. Если в колонке `Total Score` стоит положительное число, будем считать, что кошка среагировала на слово. В вектор `noun4` запишем единицу. В противном случае — ноль. 
- Вытащим по аналогии информацию о том, откликнулась ли кошка на своё имя (пометка `Test Cat's Name`) в колонке `stimuli`. Запишем её в вектор `name`.
- Используем тест о равенстве пропорций для зависимых выборок, чтобы ответить на вопрос "А правда ли кошка откликается на своё имя?" В переменную `pval` запишем $p$-значение, полученное при проверке этой гипотезы. 

<!-- grade_id: cell-4ab37a21e0f8a42c, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

noun4 = ... 
name = ... 

z_obs = ...
pval = ...


```

<!-- grade_id: cell-473051ae67ddc9a6, points: 1, grade: true, solution: false -->
```python
# проверка того, что задание решено верно
assert sum(name) == 14
assert np.abs(z_obs - 1.7888543819998317) < 1e-5

```

Выводы учёных могут отличаться от наших. Результат проверки гипотезы зависит от многих внешних факторов. Например, от громкости произнесённого кошачьего имени. Все такие факторы надо как-то изолировать.

[В оригинальном исследовании](https://www.nature.com/articles/s41598-019-40616-4) учёные реализовывали более сложные статистические процедуры. Они пытались проконтролировать все внешние факторы, которые могли бы повлиять на результаты эксперимента. О том, как это можно делать, мы с вами поговорим в будущем. 

![](https://cdn.nplus1.ru/images/2019/04/04/47a222337ec0984bcec1644f12b4c328.gif)

## Матан

Лекции по матану на экономическом факультете читают два преподавателя: Злой и Добрый. Каждый год они чередуются. Среди студентов ходит слух, что в год, когда преподаёт Злой, отчисляют куда больше народу. В таблице `matan.csv` лежит информация по $418$ студентам экономического факультета. 

В столбце `is_angry` стоит единица, если этому студенту преподавал Злой. В столбце `target_1` стоит единица, если человек закончил первый курс. Проверьте гипотезу о том, что Злой отчисляет больше народу. В переменную `pval` запишите $p$-значение, получившееся при проверке гипотезы. 

```python
df = pd.read_csv('./data/matan.csv', sep='\t')
df.head()
```

<!-- grade_id: cell-06da65b1aca523bb, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

pval = ... 
z_obs = ...


```

```python
# проверка того, что задание решено верно
assert np.abs(np.abs(z_obs) - 0.29897)  < 1e-5

```

## Бутерброд

Вася решил проверить известное утверждение о том, что бутерброд падает маслом вниз. Для этого он провёл серию из $200$ испытаний. Бутерброд упал маслом вниз $105$ раз. Маслом вверх $95$ раз. Можно ли утверждать, что бутерброд падает маслом вниз так же часто, как и маслом вверх на уровне значимости $5\%$? В качестве ответа сохраните в `pval` $p$-значение. 

__Примечание:__ для поиска дисперсии используйте смещённый вариант. Либо по формуле $\frac{p\cdot(1-p)}{n}$ либо с помощью `std(ddof=0)`. При поиске `pval` обратите внимание, что алтернатива двусторонняя. А ещё [прочитайте историю](https://absentis.livejournal.com/27153.html) про то, почему Коперник мазал бутеры маслом. 

<!-- grade_id: cell-db59ee9982bc3985, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

n = 200
m = 105

pval = ... 
z_obs = ...


```

```python
# проверка того, что задание решено верно
assert np.abs(z_obs - 0.70799) < 1e-5

```

<img src="https://i.ibb.co/Ks57rDV/student-teacher-dist.png" height="300" width="600"> 

# Часть 2 - Непараметрические гипотезы


## Однородная игра

Киллер - это весёлая лагерная игра. Краткие правила игры:

- Игроки пишут свои имена на карточках и кладут их в шляпу
- Игроки рандомно вытаскивают карточки из шляпы 
- Полученное имя - имя жертвы
- Каждый игрок одновременно и убийца и жертва
- Чтобы убить другого игрока, надо оказаться с ним наедине и застрелить из пальца
- Жертва убитого становится новой жертвой игрока
- Побеждает тот, кто совершит наибольшее число убийств

На экономическом факультете было два раунда игры в киллера. Все убийства регистрировались в гугл-форме. В табличках `killer_time1.csv` и `killer_time2.csv` лежит время в часах, которое проходило между убийствами.

```python
killer1 = pd.read_csv("./data/killer_time1.csv").hours_between_kill.values
killer2 = pd.read_csv("./data/killer_time2.csv").hours_between_kill.values
```

```python
plt.hist(killer1, bins=20);
```

```python
plt.hist(killer2, bins=20);
```

Распределение выглядит, как экспоненциальное. Более того, по гистограммам возникает ощущение, что это одинаковое распределение. С помощью теста Колмогорова-Смирнова проверьте гипотезу о том, что обе выборки пришли из одного и того же распределения, $p$-значение, полученное при проверке запишите в переменную `pvalue1`. 

<!-- grade_id: cell-f8134e5d39cba5cd, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

pvalue1 = ...

```

С помощью теста Колмогорова проверьте гипотезу о том, что время для первой игры имеет экспоненциальное распределение с параметром $\lambda = 0.026$, $p$-значение, полученное при проверке запишите в переменную `pvalue2`. 

<!-- grade_id: cell-da5a87e744840bab, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

pvalue2 = ...

```

<!-- grade_id: cell-be55b7fce581170d, points: 2, grade: true, solution: false -->
```python
# проверка того, что задание решено верно
assert pvalue1 > 0.5

```

## Мыши и зеркало

$16$ лабораторных мышей были помещены в двухкомнатные клетки, в одной из комнат висело зеркало. С целью установить, есть ли у мышей какие-то предпочтения насчет зеркал, измерялась доля времени, которое каждая мышь проводила в каждой из своих двух клеток.

![](https://ars.els-cdn.com/content/image/1-s2.0-S0168159104000061-gr1.gif)

В файле `mouses_mirror.txt` лежит информация о том, какую долю времени мышка проводила в клетке с зеркалом. 

```python
x = pd.read_csv('./data/mouses_mirror.txt', header = None)[0].values
plt.hist(x, density=True);
```

Проверьте гипотезу о том, что мышкам безразлично в какой комнате проводить время с помощью критерия знаков. Сформулируйте нулевую гипотезу, перечислите основные особенности этого критерия. Запишите в переменную `pvalue_1` получившееся в результате этого $p$-значение. 

<!-- grade_id: cell-6b6cc3cc3f9d2e2d, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

pvalue_1 = ... 

```

Проверьте гипотезу о том, что мышкам безразлично в какой комнате проводить время с помощью теста Вилкоксона. Сформулируйте нулевую гипотезу, перечислите основные особенности этого критерия.  Запишите в переменную `pvalue_2` получившееся в результате этого $p$-значение. 

<!-- grade_id: cell-689716dd8c46567c, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

pvalue_2 = ... 

```

Проверьте гипотезу о том, что мышкам безразлично в какой комнате проводить время с помощью асимптотического $z$-теста. Сформулируйте нулевую гипотезу, перечислите основные особенности этого критерия.  Запишите в переменную `pvalue_3` получившееся в результате этого $p$-значение. 

<!-- grade_id: cell-4ce505503163ae74, points: 0, grade: false, solution: true -->
```python
from statsmodels.stats.weightstats import ztest,zconfint

### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

pvalue_3 = ... 

```

<!-- grade_id: cell-2f9dfc7c1834c8a5, points: 2, grade: true, solution: false -->
```python
# проверка того, что задание решено верно
assert pvalue_1 < 0.05
assert pvalue_2 > 0.05
assert pvalue_3 > 0.05

```

Объясните, почему `pvalue` от теста к тесту ведёт себя именно так. Подробнее про опыт с мышками читайте [в оригинальном исследовании.](https://www.sciencedirect.com/science/article/abs/pii/S0168159104000061)

## Энергетик

Исследовательница Глафира считает, что любовь к энергетическим напиткам и успешность сдачи экзамена по математической статистике должны быть как-то связаны. Опросив $200$ своих однокурсников, она получила следующие результаты:

- Сдал, пьёт энергетик: 20
- Сдал, не пьёт энергетик: 120
- Завалил, пьёт энергетик: 40
- Завалил, не пьёт энергетик: 20

Правда ли, что энергетик оказывает на успешность сдачи экзамена какое-то влияние? Проверьте гипотезу с помощью критерия Пирсона. В качестве ответа запишите в переменную `pvalue` получившееся $p$-значение.

<!-- grade_id: cell-8806588bc0865379, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

pvalue = ... 

```

<!-- grade_id: cell-920f09f17b0aafde, points: 2, grade: true, solution: false -->
```python
# проверка того, что задание решено верно
assert pvalue < 0.05

```

## Бомбардировки

В таблице `london_07-09-1940.tsv` лежат данные о бомбардировке Лондона 7 сентября 1940 года немецкой авиацией. Интересно выяснить, насколько прицельно велась бомбардировка. Если немецкие войска знали, где именно находятся британские военные объекты, они вели бы бомбардировку пристрельно.  

Ещё больше данных о бомбардировке Лондона ищите здесь: http://bombsight.org/data/sources/

```python
df = pd.read_csv('./data/london_07-09-1940.tsv', sep='\t')
print(df.shape)
df.head()
```

```python
plt.scatter(df.x, df.y);
```

В колонках $x$ и $y$ лежат координаты точки, куда упала бомба. Разбейте Лондон решёткой размера $30 \times 30$ на $900$ небольших квадратных участков. С помощью функции `binned_statistic_2d` из пакета `scipy.stats`, подсчитайте как часто снаряд попадал на какой из участков. Поместите информацию об этом в переменную `ret`. 

<!-- grade_id: cell-2a06adb25a0b5542, points: 0, grade: false, solution: true -->
```python
binx = np.linspace(df.x.min(), df.x.max(), 31)
biny = np.linspace(df.y.min(), df.y.max(), 31)

### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

ret = ... 

sns.heatmap(ret.statistic);
```

Распакуйте информацию о попаданиях в участки в вектор `z`. Нарисуте гистограмму для получившегося распределения. 

<!-- grade_id: cell-1b692a90383e2eff, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

z = ...


```

Участков много, а вероятность попадания на отдельный участок мала. При справедливости гипотезы о низкой точности стрельбы мы можем воспользоваться законом редких событий:

> Если в схеме испытаний Бернулли $p \to 0$, а $n \to \infty$, оно сходится к распределению Пуассона. 

Получается, что __при верности гипотезы о низкой точности стрельбы,__ число попаданий в квадрат имеет распределение Пуассона. Дополнительно будем предполагать, что попадания на разные участки независимы друг от друга. 

C помощью критерия Пирсона проверьте гипотезу о том, что число попаданий на участок имеет распределение Пуассона. Прежде, чем вычислять статистику, объедините все наблюдения из хвоста в одно по аналогии с лекцией. Хвостом будем считать все наблюдения $\ge 9$. Получившуюся в результате расчётов статистику запишите в переменную `T`.

<!-- grade_id: cell-ae799b5490316110, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# your code below

T = ...

```

<!-- grade_id: cell-2986a877274d9a7e, points: 2, grade: true, solution: false -->
```python
# проверка того, что задание решено верно
assert T > 6000

```

Какой вывод мы можем из этого сделать? Правда ли, что бомбардировка осуществлялась прицельно? 

# Часть 3 - Самостоятельный анализ


- Ниже приведены несколько датасетов, их необходимо самостоятельно проанализировать и ответить на определенные вопросы
- Для ответа на эти вопросы нужно сделать следующее
    - Сделать хорошую визуализацию исследуемых переменных
    - Для каждой из задач подобрать **все подходящие** параметрические и непараметрические критерии, проанализировать результаты этих критериев (тут имеет смысл обернуть все критерии в фунцию или словарь - как удобно)
- Для каждого набора данных проверьте нормальность выборки и допустимость использования параметрических критериев (можно прокомментировать гистограмму и воспользоваться соответствующими критериями согласия)
- Не забудьте оставить краткие комментарии к вашим действиям и сформулировать выводы, которые вы сделали в ходе анализа


## Анорексия

В файле anorexia.txt записан вес пациентов до начала терапии анорексии и после ее окончания. Была ли терапия в эффективна?


```python
df = pd.read_csv('./data/anorexia.txt', sep='\t')
df.head()

```

Output:
```
Before  After
0    80.5   82.2
1    84.9   85.6
2    81.5   81.4
3    82.6   81.9
4    79.9   76.4
```

## Метилфенидат
В рамках исследования эффективности препарата метилфенидат пациенты с синдромом дефицита внимания и гиперактивности в течение недели принимали либо метилфенидат, либо плацебо. В конце недели каждый пациент проходили тест на способность к подавлению импульсивных поведенческих реакций. На втором этапе плацебо и препарат менялись, и после недельного курса каждый испытуемые проходили второй тест. Был ли эффект от применения препарата? Данные находятся в файле methylphenidate.txt.



```python
df = pd.read_csv('./data/methylphenidate.txt', sep=' ')
df.head()

```

Output:
```
Placebo  Methylphenidate
0       57               62
1       27               49
2       32               30
3       31               34
4       34               38
```

```python

```

## Сиэтл
В файле seattle.txt записаны цены на недвижимость в одном из районов Сиэтла в 2001 году и в 2002 году (объекты недвижимости выбирались случайно). Изменились ли в среднем цены в этом районе за год?


```python
df = pd.read_csv('./data/seattle.txt', sep='\t')
df.head()

```

Output:
```
Price  Year
0  142.0  2002
1  232.0  2002
2  132.5  2002
3  200.0  2002
4  362.0  2002
```

## Расход топлива
В файле mtcars.csv находятся данные из американского журнала Motor Trend 1974 года. Они описывают расход топлива в зависимости от 10 характеристик автомобиля (все автомобили 1973-1974 года). А именно, в данных вы найдете следующие столбцы:
1. mpg — расход топлива (миль/галлон); 
2. cyl — количество цилиндров мотора; 3) disp — объем мотор (куб. дюйм);
4. hp — мощность мотора (в лошадиных силах);
5. drat — передаточное число редуктора заднего моста;
6. wt — вес (тысяч фунтов);
7. qsec — время, за которое автомобиль проезжает 1/4 мили;
8. vs — тип двигателя (0 = V-образный, 1 = рядный);
9. am — тип коробки передач (0 = автоматическая, 1 = ручная);
10. gear — количество передних передач; 11) carb — количество карбюраторов.



- Влияет ли тип двигателя на расход топлива? 
- А тип коробки передач?
- Взаимосвязаны ли тип двигателя и тип коробки передач?

```python
df = pd.read_csv('./data/mtcars.csv', index_col=0)
df.head()

```

Output:
```
mpg  cyl   disp   hp  drat     wt   qsec  vs  am  gear  \
Mazda RX4          21.0    6  160.0  110  3.90  2.620  16.46   0   1     4   
Mazda RX4 Wag      21.0    6  160.0  110  3.90  2.875  17.02   0   1     4   
Datsun 710         22.8    4  108.0   93  3.85  2.320  18.61   1   1     4   
Hornet 4 Drive     21.4    6  258.0  110  3.08  3.215  19.44   1   0     3   
Hornet Sportabout  18.7    8  360.0  175  3.15  3.440  17.02   0   0     3   

                   carb  
Mazda RX4             4  
Mazda RX4 Wag         4  
Datsun 710            1  
Hornet 4 Drive        1  
Hornet Sportabout     2
```

#### vs, mpg

```python

```

#### am, mpg

```python

```

#### vs, am

```python

```

```python

```